```python
version = "REPLACE_PACKAGE_VERSION"
```

---
# Assignment 1 Part 1: N-gram Language Models (40 pts)

In this assignment, we're going to train an n-gram language model that is able to "imitate" William Shakespeare's writing. 

```python
# Configure nltk

import nltk

nltk_data_path = "assets/nltk_data"
if nltk_data_path not in nltk.data.path:
    nltk.data.path.append(nltk_data_path)
```

## Question 1: Load the dataset (10 pts)

As the first step towards imitating Shakespeare's writing, you will create a function called `load_data` that loads his original *Sonnets* from `assets/gutenberg/THE_SONNETS.txt`. This function should accomplish the following: 

* **Extract sentences from the data file.** Of course, depending on the nature of the task at hand, what constitutes a *sentence* can vary. In the context of this assignment, we will define a sentence as just a line of a sonnet, regardless of the punctuation at end. In addition, we will ignore the boundaries of the sonnets --- that is, we are not dealing with $154$ individual *sonnets* but rather $154 \times 14 = 2156$ *sentences* (actually only $2155$ sentences, as *Sonnet 99* has [15 lines](https://www.google.com/search?ei=po4hX9jzN4rKtQaL7K-wDg&q=why+is+sonnet+99+15+lines&oq=why+is+sonnet+99+15+lines&gs_lcp=CgZwc3ktYWIQAzoECAAQR1ChGlihGmCqHGgAcAJ4AIABXIgBXJIBATGYAQCgAQGqAQdnd3Mtd2l6wAEB&sclient=psy-ab&ved=0ahUKEwjY3snX3PLqAhUKZc0KHQv2C-YQ4dUDCAw&uact=5) but *Sonnet 126* has only [12](https://www.google.com/search?ei=d4whX6_uH9a4tQbUiISoDA&q=why+is+sonnet+126+only+12+lines&oq=o+thou+my+lovely+boy+who+in+thy+power&gs_lcp=CgZwc3ktYWIQARgBMgQIABBHMgQIABBHMgQIABBHMgQIABBHMgQIABBHMgQIABBHMgQIABBHMgQIABBHUABYAGCWEWgAcAJ4AIABAIgBAJIBAJgBAKoBB2d3cy13aXrAAQE&sclient=psy-ab)). We encourage you to explore alternative definitions of a sentence on your own; for example, an entire sonnet could be modelled as a sentence. Finally, make sure that the newline character `\n` at the end of each line is dropped. 


* **Tokenise each extracted sentence.** While it's ambiguous what a sentence is, what constitutes a "*word*" is even more task-dependent. Do punctuations count as "words"? Are two "words" with the same spelling but different casing considered identical? Since what a text file contains is nothing more than a squence of characters, there are many possible ways of grouping these characters to form "words" that are subsequently taken as input by a program. To distinguish what's actually taken as input by a program from a linguistic word, we call the former a *token*. The process of producing a list of tokens out of a sentence is then called *tokenisation*. At this step, you will first lower-case each sentence extracted from the previous step entirely and then tokenise each lower-cased sentence. You may use any tokeniser of your choice, such as `word_tokenize` from the `nltk` library. The grading of the assignment doesn't depend on your choice of the tokeniser. 

**This function should return a `list` of length 2155, where each element is a `list` of `str` representing the tokens of each sentence as produced by the tokeniser of your choice. An example output would be:**

```
[['from', 'fairest', 'creatures', 'we', 'desire', 'increase', ','],
 ['that', 'thereby', 'beauty', '’', 's', 'rose', 'might', 'never', 'die', ','],
 ...
 ['came', 'there', 'for', 'cure', 'and', 'this', 'by', 'that', 'i', 'prove', ','], 
 ['love', '’', 's', 'fire', 'heats', 'water', ',', 'water', 'cools', 'not', 'love', '.']]
```

<!-- grade_id: cell-876ee5c9fdc4f347, points: 0, grade: false, solution: true -->
```python
def load_data():
    """
    Load text data from a file and produce a list of token lists
    """
    
    sentences = []
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return sentences
```

<!-- grade_id: cell-436f8cfe2405ffe8, points: 10, grade: true, solution: false -->
```python
# Autograder test

stu_ans = load_data()

assert isinstance(stu_ans, list), "Q1: Your function should return a list. "
assert len(stu_ans) == 2155, "Q1: There should be 2155 sentences. "

for index, tokens in enumerate(stu_ans):
    assert isinstance(tokens, list), f"Q1: The element at index {index} of your answer list should be a list. "
    for token in tokens:
        if token.isalpha():
            assert token.islower(), f'Q1: Token "{token}" in the sentence at index {index} is not lower-cased. '
        
        assert token != "\n", f'Q1: You should drop the "\\n" character in the sentence at index {index}. '

del stu_ans
```

Output:
```
Cell In[4], line 3
      1 # Autograder test
----> 3 stu_ans = load_data()
      5 assert isinstance(stu_ans, list), "Q1: Your function should return a list. "
      6 assert len(stu_ans) == 2155, "Q1: There should be 2155 sentences. "

Cell In[3], line 9, in load_data()
      6 sentences = []
      8 # YOUR CODE HERE
----> 9 raise NotImplementedError()
     11 return sentences

NotImplementedError:
```

## Question 2: Build vocabulary (15 pts)

Next, we need a "vocabulary" that contains all the unique tokens. Moreover, as mentioned in the lecture, we often pad a sentence with `<s>` and `</s>` to indicate its start and end when working with n-gram language models; therefore, these two special tokens should also be included in our vocabulary. Complete the function below to build a vocabulary. The order in which the tokens are stored doesn't matter. 

**This function should return a `list` of unique tokens, including `<s>` and `</s>`. An example output would be:**

```
['refuse', 'enjoyed', ..., '<s>', '</s>']
```

<!-- grade_id: cell-be2a75b95ae4c066, points: 0, grade: false, solution: true -->
```python
def build_vocab(sentences):
    """
    Take a list of sentences and return a vocab
    """
    
    vocab = []
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return vocab
```

<!-- grade_id: cell-bd592f77659516e2, points: 15, grade: true, solution: false -->
```python
# Autograder tests

stu_sents = load_data()
stu_vocab = build_vocab(stu_sents)

assert isinstance(stu_vocab, list), "Q2: Your function should return a list. "
assert stu_vocab, "Q2: Your vocab is empty. "
assert "<s>" in stu_vocab, "Q2: Remember to include special token <s>. "
assert "</s>" in stu_vocab, "Q2: Remember to include special token </s>. "
assert len(set(stu_vocab)) == len(stu_vocab), "Q2: Your vocab contains duplicated tokens. "

# Some hidden tests

del stu_sents, stu_vocab
```

## 

## Question 3: Generate all $n$-grams (15 pts)

Now let's write a function to generate all $n$-grams for each sentence. This can be accomplished in two steps:
* **Pad each sentence with `<s>` and `</s>` for $n \geq 2$**. You need $n - 1$ paddings on both ends of a sentence, so that there are two $n$-grams that model the first and the last token, respectively. You may implement the padding function yourself or use the `pad_both_ends` function from the `nltk` library. 


* **Generate $n$-grams on the padded sentences**. Check out the `ngrams` function from `nltk`. For a sentence with $\ell$ tokens excluding paddings, the maximum number of possible $n$-grams generated from its padded version should be $\ell + n - 1$. Think about why. 

**Complete the function below to return a `list`, where each element of the `list` is a either a list or a "generator object" produced by the `ngrams` function, representing a sequence of all $n$-grams generated from each appropriately padded sentence. If the argument `n=2`, the autograder would accept either of the example outputs below:**

```
[<generator object ngrams at 0x7f77ed778b10>,
 <generator object ngrams at 0x7f77e7d934f8>,
 ...
 <generator object ngrams at 0x7f77e7d751b0>,
 <generator object ngrams at 0x7f77e7d75228>]
```

**OR**

```
[
    [('<s>', 'from'), ('from', 'fairest'), ('fairest', 'creatures'), ('creatures', 'we'), ('we', 'desire'),
     ('desire', 'increase'), ('increase', ','), (',', '</s>')], 
    
    [('<s>', 'that'), ('that', 'thereby'), ('thereby', 'beauty'), ('beauty', '’'), ('’', 's'), ('s', 'rose'), 
     ('rose', 'might'), ('might', 'never'), ('never', 'die'), ('die', ','), (',', '</s>')],
     
    ...
    
    [('<s>', 'came'), ('came', 'there'), ('there', 'for'), ('for', 'cure'), ('cure', 'and'), ('and', 'this'), 
     ('this', 'by'), ('by', 'that'), ('that', 'i'), ('i', 'prove'), ('prove', ','), (',', '</s>')], 
    
    [('<s>', 'love'), ('love', '’'), ('’', 's'), ('s', 'fire'), ('fire', 'heats'), ('heats', 'water'), 
     ('water', ','), (',', 'water'), ('water', 'cools'), ('cools', 'not'), ('not', 'love'), ('love', '.'), 
     ('.', '</s>')]
]
```

<!-- grade_id: cell-79211162e032488b, points: 0, grade: false, solution: true -->
```python
def build_ngrams(n, sentences):
    """
    Take a list of unpadded sentences and create all n-grams as specified by the argument "n" for each sentence
    """
    
    all_ngrams = []
    #HINT: use pad_both_ends from nltk library   
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return all_ngrams
```

<!-- grade_id: cell-66bf33f44f8b83bf, points: 15, grade: true, solution: false -->
```python
# Autograder tests

stu_n = 4
stu_sents = load_data()
old_hash = hash(tuple([tuple(sent) for sent in stu_sents]))
stu_ngrams = build_ngrams(stu_n, stu_sents)

assert isinstance(stu_ngrams, list), "Q3: Your function should return a list. "
assert stu_ngrams, "Q3: Your ngrams list is empty. "

# Check that your function does not modify 'stu_sents' in place
new_hash = hash(tuple([tuple(sent) for sent in stu_sents]))
assert new_hash == old_hash, "Q3: Your function should not modify its argument 'sentences' in place. "

# Some hidden tests


del stu_n, stu_sents, stu_ngrams
```

Now that we have completed all the preparation work for imitating William Shakespeare's writing, it's time to take a break. We will resume in Assignment 1 Part 2 to finish training an $n$-gram language model. See you there!

```python

```